```python
# Initialize Otter
import otter
grader = otter.Notebook("env-validation.ipynb")
```

# Environment Validation

This is a sample assignment notebook for validating that Otter works in your environment.

To test your environment, run

```console
otter assign env-validation.ipynb dist -v
```

You should see a directory called `dist` created with the autograder and student subdirectories and a `data.csv` file in each. You should also see an autograder zip file in the autograder directory. You should also see in the output of the `otter assign` command that the tests were run on the notebook and they succeded.

```python
import pandas as pd
import random
```

**Question 1:** Fill in the `hailstone` function below which returns the hailstone sequence of the number `n`.

```python
def hailstone(n):
    ...

hailstone(7)
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2.** Evaluate the integral:

$$
\int x e^x \, \mathrm dx
$$

$$
\begin{aligned}
\int x e^x \, \mathrm dx &= \texttt{YOUR MATH HERE} \\
\end{aligned}
$$

<!-- END QUESTION -->

**Question 3:** Load the data from `data.csv` and perform an A/B test to determine whether the population in condition `A` has a statistically significant difference in `val` rfom the population in condition `B`. 

```python
df = pd.read_csv("data.csv")

def test_stat(df, cond_col):
    """
    Calculate the test state on a dataframe. Assumes that the 
    """
    return abs(df[df[cond_col] == "A"]["val"].mean() - df[df[cond_col] == "B"]["val"].mean())

observed = ...

stats = []
for _ in range(1000):
    df["shuffled_cond"] = ...
    ...

p_value = ...
print(f"A/B test p-value: {p_value}")
```

```python
grader.check("q3")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert hailstone(1) == [1]
```
```
>>> assert hailstone(2) == [2, 1]
```
```
>>> assert hailstone(3) == [3, 10, 5, 16, 8, 4, 2, 1]
```


Tests for q3 (points: None)
```
>>> assert len(stats) == 1000
```
```
>>> assert 0 <= p_value <= 1
```
